Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Guilherme Pereira"
COLLABORATORS = "Peaceful Piano"
```

---

# Lista 04 - Teste de Hipótese

```python
#Defina a semente de geração de números aleatórios
#Não gerem mais números aleatórios do que o necessário, ou a correção dará errado
import numpy as np
import matplotlib.pyplot as plt
```

```python
#Fixa a semente geradora de numeros aleatorios para gerar os mesmos resultados a cada execucao
np.random.seed(42)
```

# Exercício 01:

Em uma determinada universidade, foi realizada uma entrevista em que se desejava coletar respostas de alunos de todas as grandes áreas de conhecimento. Dentre os 1000 entrevistados, cada área teve a seguinte quantidade de alunos:

* Exatas: 100
* Humanas: 600
* Biológicas: 300

No entanto, a qualidade da amostra foi questionada. Argumentou-se que a proporção de alunos de Humanas entrevistados foi muito maior que a proporção de alunos da mesma área na universidade como um todo. Na universidade, há cadastro de:

* 2000 alunos de Exatas
* 8000 alunos de Humanas
* 2500 alunos de Biológicas

A) Retorne uma lista de True ou False de acordo com as afirmativas abaixo:

    1: A respeito da hipótese nula, podemos realizar testes que nos permitam aceitá-la como verdadeira.
    2: Podemos definir a seguinte hipótese nula: "A amostra utilizada na entrevista não foi gerada de forma uniformemente aleatória".
    3: Podemos definir a seguinte hipótese nula: "A amostra utilizada na entrevista foi gerada de forma uniformemente aleatória".
    
Exemplo:

```python
def q1a():
    return [ False, False, False ]
```

<!-- grade_id: cell-8df9138889cc34a2, points: 0, grade: false, solution: true -->
```python
def q1a():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-921d6c90865df759, points: 1, grade: true, solution: false -->
```python

```

B) Para comparar as distribuições, implemente a métrica de Total Variation Distance (TVD):

$$ \mathit{TVD}(p, q) = \frac{1}{2}|p - q| = \frac{1}{2}\sum_{i=1}^n |p_i - q_i| $$    
onde $p$ e $q$ são os vetores de proporções

<!-- grade_id: cell-3001a387e9432443, points: 0, grade: false, solution: true -->
```python
def tvd(p, q):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6d02be67dbb94b79, points: 1, grade: true, solution: false -->
```python
from numpy.testing import *
pop = np.array([0.15, 0.18, 0.12, 0.54, 0.01])
samp = np.array([0.26, 0.08, 0.08, 0.54, 0.04])
assert_equal( tvd(pop,samp), 0.14 )
```

<!-- grade_id: cell-16f566cad8421eec, points: 1, grade: true, solution: false -->
```python

```

C) Calcule o TVD da amostra utilizada na entrevista

__Dica:__ A Total Variation Distance é uma métrica de distância entre distribuições de probabilidade, não entre quantidades absolutas.

<!-- grade_id: cell-ccb92b26dd65a82d, points: 0, grade: false, solution: true -->
```python
def q1c():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-28fc937dbf0d704c, points: 1, grade: true, solution: false -->
```python

```

D) Implemente a função que simule uma amostra uniformemente aleatória, retornando uma lista de proporções para cada categoria (Exatas, Humanas, Biológicas).

Exemplo: 
```python
def gera_amostra(tamanho_amostra):
    return [0.6, 0.0, 0.4]
```
__Dica:__ crie uma lista para a população (indicando a área de conhecimento do estudante por 0,1 ou 2, por exemplo) e use a função ```numpy.random.shuffle``` para embaralhar os dados. Depois, pegue os primeiros elementos da lista para formar a amostra, e calcule as proporções a partir disso.

Por exemplo, o vetor [0, 0, 1, 1, 1, 2] representa uma população de tamanho 6, com dois elementos na classe 0, três elementos na classe 1 e um elemento na classe 2. As proporções são 2/6, 3/6 e 1/6, respectivamente.

<!-- grade_id: cell-4e7e282297db6488, points: 0, grade: false, solution: true -->
```python
def calcula_proporcao(amostra, valor):
    return sum(1 for x in amostra if x == valor) / len(amostra)

def gera_amostra(tamanho_amostra):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-50852e242dfddbb0, points: 1, grade: true, solution: false -->
```python

```

E) Gere 10000 amostras, e plote o gráfico de distribuição dos valores de TVD. Coloque no gráfico um ponto mostrando o TVD da amostra da entrevista.

<!-- grade_id: cell-435a898fe2b6c63d, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

F) Escreva uma função q1f(x) que retorne a partir de qual valor de TVD não estamos mais dentro do intervalo de x% de confiança.

__Dica__: Utilize o método baseado em boostrap, ou seja, calcule o TVD para um número grande de amostras da população, como foi feito no passo anterior. Em seguida, encontre o percentil dos TVDs calculados que corresponde ao nível de confiança de 90%.

Retorne apenas o lado direito do intervalo de confiança para o TVD.


<!-- grade_id: cell-9a5e069bc5ba5d72, points: 0, grade: false, solution: true -->
```python
def q1f(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-be1808c9bbd14418, points: 0, grade: true, solution: false -->
```python
from numpy.testing import *
assert_almost_equal( q1f(95), 0.034, 3 )
assert_almost_equal( q1f(90), 0.031, 3 )
assert_almost_equal( q1f(80), 0.026, 3 )
```

<!-- grade_id: cell-b5ce0fa1e395ffa3, points: 1, grade: true, solution: false -->
```python

```

# Exercício 02:

Usando os dados do arquivo ```enem2015.csv```, vamos realizar um teste por permutação.

```python
import pandas as pd
```

```python
df = pd.read_csv('enem2015.csv')
df.head()
```

A) Implemente uma função que retorne a diferença entre a média das notas das escolas públicas federais e a média das notas de  escolas privadas. Exemplo:

```python
def q2a(df):
    media_federal = 631.0
    media_particular = 620.0
    return 11.0
```

<!-- grade_id: cell-72b78cf05ad691a9, points: 0, grade: false, solution: true -->
```python
def q2a(df):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-be0bf46c5dc1853d, points: 1, grade: true, solution: false -->
```python

```

B) Realize permutações com os rótulos (Federal / Privada), de tamanhos ```[10,100,500,1000,5000]```. Plote gráficos mostrando a distribuição das diferenças entre as médias notas, conforme se aumenta a quantidade de permutações. Ou seja, um gráfico para 10 permutações, um para 100, e assim por diante.

Como exemplo, a imagem abaixo mostra o gráfico para 100 permutações.

```python
from IPython.display import Image
Image('100permutacoes.png')
```

<!-- grade_id: cell-a4244ab97beaa278, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

B) Escreva abaixo o que você observou com os gráficos.

<!-- grade_id: cell-79def4591cae6a4d, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE